```
BEGIN ASSIGNMENT
requirements: requirements.txt
init_cell: false
export_cell: false
template_pdf: false
files:
    - train.csv
generate:
    token: ''
    course_id: 62982        # required
    assignment_id: 9    # required
    filtering: true # true is the default
    seed: 42
    show_stdout: true
    show_hidden: true
    show_all_public: true
    visibility: visible
```

# Exercise sheet 9
**Hello everyone!**

**Points: 15**

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [Piazza](https://piazza.com/class/leibdr4lk8n3w4), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids23).

Please submit this exercise sheet on **GRADESCOPE**.
Naming conventions:
* Name of notebook: "09-exercise-pids2023.ipynb"

Your notebook will be automatically graded using gradescope. To get graded, you need to upload the file *09-exercise-pids2023.ipynb* to gradescope. If you are using JupyterHub, make sure you have saved the current notebook, then you can download this file by following: *File*->*Download* and save the file on you computer.

The sheet is released before the exercise session on Tuesday. You have until the next exercise session on Tuesday to submit this sheet. Most of the questions of the exercise sheet can be handle during the session.

**Handout date**: 02.05.2023 <br>
**Submission date**: 09.05.2023 <br>

### Please name the variables as suggested! Be careful with the type! Otherwise the grading system will fail you!

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import numpy as np
from scipy import stats
from scipy.stats import norm
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

```python
import otter
grader = otter.Notebook()
```

# More uncertainty and confidence intervals!
Throughout this notebook, we will steadily increase our understanding of confidence intervals. While there are more mathematical notations than usual, they are included solely to enhance the clarity of the exercises. You won't need to run difficult computations other than empirical mean and empirical standard deviation.


# Question 1: Titanic dataset
We are working with a dataset containing information about passengers of the sadly famous Titanic. We will see how we can use tools from the lecture to get a sense of what factors are important to increase our probability of survival during a shipwreck.<br>

In the first part, we will load the dataset, and plot a few values to see which parameters might be important. In a second part, we will focus on one particular feature and see if it is trustworthy.

### 1a (1 point): 
We first load the dataset. If you are working locally, you should be able to load the dataset from the csv file named 'train.csv', in the same location as this current file.

```
BEGIN QUESTION
name: Question 1a        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

```python
data_source = "https://raw.githubusercontent.com/paulhendricks/titanic/master/inst/data-raw/train.csv"
data = pd.read_csv(data_source) # SOLUTION
data.head()
```

```python
## test ##
assert_equal(len(data),891)
```

### 1b (1 point):
Create a variable ***data_class_survived*** and ***data_class_not_survived*** that contains only the passengers that survived and the passagengers that didn't survive, respectively. Here, we will only keep the *Pclass* and *Survived* columns from the dataset. 

```
BEGIN QUESTION
name: Question 1b        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

```python
data_class = data[['Survived','Pclass']]
data_class_survived = data_class[data_class['Survived']==1] #SOLUTION
data_class_not_survived = data_class[data_class['Survived']==0] #SOLUTION
```

```python
## test ##
assert_true(len(data_class_survived), 342)
```

```python
## test ##
assert_true(len(data_class_survived), 549)
```

### 1c (1 point): 
Make a normalized histogram showing the class (column 'Pclass') of the passengers that survived.<br>
On the same plot, display the same histogram for the passengers that did not survive.<br>
Hint: you can use plt.hist and the density parameter.

```
BEGIN QUESTION
name: Question 1c        # (required) the path to a requirements.txt file
manual: true     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-ede3d355a0957fbe, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
_ = plt.hist(data_class_survived['Pclass'],bins=16,alpha=0.5,label='survived',density=True)
_ = plt.hist(data_class_not_survived['Pclass'],bins=16,alpha=0.5,label='not survived',density=True)
_= plt.legend()
# END SOLUTION
```

### 1d (1 point):
Similar to 1c), display the 'Fare' of passengers survived in a histogram.<br>
On the same plot, display the same histogram for the passengers that did not survive. <br>
In order to do so, create data frames for the information in column 'Fare' as you did in 1b.


```
BEGIN QUESTION
name: Question 1d        # (required) the path to a requirements.txt file
manual: true     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-f95caf2cb1886049, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
data_Fare = data[['Survived','Fare']]
data_Fare_survived = data_Fare[data_Fare['Survived']==1]
data_Fare_not_survived = data_Fare[data_Fare['Survived']==0]

plt.hist(data_Fare_survived['Fare'],bins=16,alpha=0.5,label='survived',density=True)
plt.hist(data_Fare_not_survived['Fare'],bins=16,alpha=0.5,label='not survived',density=True)
plt.legend()
# END SOLUTION
```

### 1e (2 points):
Create variables ***data_age_survived*** and ***data_age_not_survived*** to extract passengers that survived and passengers that did not.<br>
Among the passengers that survived, display ages of the passengers in a histogram. Make sure that the histogram is *normalized*<br>
On the same plot, display the same normalized histogram for the passengers that didn't survive.<br>
Hint: you can use plt.hist.

```
BEGIN QUESTION
name: Question 1e        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-0da4e66b7921cec4, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
data_age = data[['Survived','Age']] 
# END SOLUTION
data_age_survived = data_age[data_age['Survived']==1] # SOLUTION
data_age_not_survived = data_age[data_age['Survived']==0] # SOLUTION

# BEGIN SOLUTION
plt.hist(data_age_survived['Age'],bins=16,alpha=0.5,label='survived',density=True)
plt.hist(data_age_not_survived['Age'],bins=16,alpha=0.5,label='not survived',density=True)
plt.legend()
# END SOLUTION
```

```python
## test ##
assert_equal(len(data_age_survived),342)
```

```python
## test ##
assert_equal(len(data_age_not_survived),549)
```

### 1f (2 points): 
Display on the same plot, the column 'Survived' for the female and male passengers. <br>
Hint: you can use plt.hist.<br>
Among the passengers, compute the proportion of survivors that are women and store it in ***surv_female***.<br>
Among the passengers, compute the proportion of survivors that are men and store it in ***surv_male***.

```
BEGIN QUESTION
name: Question 1f        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-a66ee7dc2a8b5ed1, points: 0, grade: false, solution: true -->
```python
surv_female = 0 # SOLUTION
surv_male = 0 # SOLUTION
# BEGIN SOLUTION
data_sex = data[['Survived','Sex']]
data_female = data_sex[data_sex['Sex']=='female']
data_male = data_sex[data_sex['Sex']=='male']

plt.hist(data_female['Survived'],alpha=0.5,label='female',density=True)
plt.hist(data_male['Survived'],alpha=0.5,label='male',density=True)
plt.legend()

surv_female = (data_female['Survived']==1).sum()/len(data_sex)
surv_male = (data_male['Survived']==1).sum()/len(data_sex)
# END SOLUTION
```

```python
## test ##
assert_almost_equal(surv_female,0.2615039281705948 , places=2)
```

```python
## test ##
assert_almost_equal(surv_male,0.122334455667789 , places=2)
```

## Will your sex increase or decrease your probability of survival?

Let $(X_i)_{i\leq N_X} \overset{i.i.d.}{\sim} \mathrm{Ber}(p_X)$ denote $N_X$ Bernoulli random variables that model the outcome of a female passenger on the Titanic. <br>
Let $(Y_i)_{i\leq N_Y} \overset{i.i.d.}{\sim} \mathrm{Ber}(p_Y)$ denote $N_Y$ Bernoulli random variables that model the outcome of a male passenger on the Titanic.<br>

We aim at estimating the probailty of survival $p_X$ and $p_Y$.

Ideally we would like to get the values $p_X = \mathbb{E}[X_i]$ and $p_Y = \mathbb{E}[Y_i]$. Unfortunately, we only have access to the $N_X$ realizations $(X_i)$ and $N_Y$ realization $(Y_i)$. <br>

### 1g (1 point):
Use the law of large numbers to create an estimate $\hat p_X$ of $p_X$ and an estimate $\hat p_Y$ of $p_Y$. Name these estimators ***pX_hat*** and ***pY_hat***.

```
BEGIN QUESTION
name: Question 1g        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-f870da6474590bea, points: 0, grade: false, solution: true -->
```python
pX_hat = (data_female['Survived']==1).mean() # SOLUTION
pY_hat = (data_male['Survived']==1).mean() # SOLUTION
print(pX_hat)
print(pY_hat)
```

```python
## test ##
assert_almost_equal(pX_hat , 0.7420382165605095)
```

```python
## test ##
assert_almost_equal(pY_hat , 0.18890814558058924 , places=2)
```

Let's do some math.<br>
We have some estimate for $p_X$ and $p_Y$, but we would like to know if we can trust them. <br>

Let's first focus on $p_X$. <br>
We want to compute the probability $p_{trust X}$ that we can trust our estimate for $p_X$:
$$ p_{trust X} = \mathbb{P}\left( p_X-0.05 \leq \hat X \leq p_X + 0.05 \right). $$

Based the central limit theorem, we'll make the assumption that $\hat p_X \sim \mathcal{N}(p_X, \sigma^2/N_X)$.<br>
Then, following slides of week 8, we have
$$ p_{trust X} = \mathbb{P}\left( Z \leq 0.05\sqrt{ \frac{N_X}{p_X(1-p_X)}} \right)-\mathbb{P}\left( Z \leq -0.05\sqrt{ \frac{N_X}{p_X(1-p_X)}} \right), $$
where $Z\sim \mathcal{N}(0,1)$. <br>
As mentioned during the lecture, we don't know $p_X(1-p_X)$. Let's replace this term by $\hat p_X(1-\hat p_X)$.


### 1h (2 points):

Compute the value 
$$ p1 = \mathbb{P}\left( Z \leq 0.05\sqrt{ \frac{N_X}{\hat p_X(1-\hat p_X)}} \right)$$
and 
$$ p2 = \mathbb{P}\left( Z \leq -0.05\sqrt{ \frac{N_X}{\hat p_X(1-\hat p_X)}} \right).$$
Store these values in variables named $p1$ and $p2$.<br>
Hint: you can have access to the cumulative distribvutuion function (cdf) of a Gaussian random variable by using the function *norm* (for normal) from the library *scipy.stats* (already imported).

```
BEGIN QUESTION
name: Question 1h        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-11532a187e7a16cc, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
N_X = len(data_female)
# END SOLUTION
p1 = norm.cdf(0.05*np.sqrt(N_X/(pX_hat*(1-pX_hat)))) # SOLUTION
p2 = norm.cdf(-0.05*np.sqrt(N_X/(pX_hat*(1-pX_hat)))) # SOLUTION
print(p1, p2)
```

```python
## test ##
data_sex = data[['Survived','Sex']]
data_female = data_sex[data_sex['Sex']=='female']
data_male = data_sex[data_sex['Sex']=='male']
N_X = len(data_female)
assert_almost_equal(p1 , norm.cdf(0.05*np.sqrt(N_X/(pX_hat*(1-pX_hat)))) , places=2)
```

```python
## test ##
data_sex = data[['Survived','Sex']]
data_female = data_sex[data_sex['Sex']=='female']
data_male = data_sex[data_sex['Sex']=='male']
N_X = len(data_female)
assert_almost_equal(p2 , norm.cdf(-0.05*np.sqrt(N_X/(pX_hat*(1-pX_hat))))  , places=2)
```

### 1i (1 point): 
Based on the previous assumptions, what is the probability that $\hat p_X \in [p_X - 0.05, p_X+0.05]$? Store your answer in a variable named ***proba_trust_X***.

```
BEGIN QUESTION
name: Question 1i        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-c5f50be87fabc3a0, points: 0, grade: false, solution: true -->
```python
proba_trust_X = p1 - p2 # SOLUTION
print(proba_trust_X)
```

```python
## test ##
assert_almost_equal(proba_trust_X , p1 - p2 , places=2)
```

### 1j (1 point): 
Using a similar approach, compute the probability that $\hat p_Y \in [p_Y - 0.05, p_Y+0.05]$? Store your answer in a variable named ***proba_trust_Y***.

```
BEGIN QUESTION
name: Question 1j        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-b6aeee80244a9dd3, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
N_Y = len(data_male)
p1 = norm.cdf(0.05*np.sqrt(N_Y/(pY_hat*(1-pY_hat)))) 
p2 = norm.cdf(-0.05*np.sqrt(N_Y/(pY_hat*(1-pY_hat)))) 
# END SOLUTION
proba_trust_Y = p1 - p2 # SOLUTION
print(proba_trust_Y)
```

<!-- grade_id: cell-b361b39476a7faf5, points: 4, grade: true, solution: false -->
```python
## test ##
data_sex = data[['Survived','Sex']]
data_female = data_sex[data_sex['Sex']=='female']
data_male = data_sex[data_sex['Sex']=='male']
N_Y = len(data_male)
assert_almost_equal(proba_trust_Y , norm.cdf(0.05*np.sqrt(N_Y/(pY_hat*(1-pY_hat))))- norm.cdf(-0.05*np.sqrt(N_Y/(pY_hat*(1-pY_hat))))  , places=2)
```

# What about the children?
Should we remove the children in the previous analysis about the sex? If we do so, the number of women drop from 314 to 259. And we have seen in the lecture that the lower the number of sample, the worse is the confidence interval. 
<br>
But anyway, let's finish by checking that children have a very high probability of survival.

### 1k (2 points): 
Compute the probability of survival of a child (age<18) and store it in ***proba_surv_child***.

```
BEGIN QUESTION
name: Question 1k        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-f85e1d8e3c4f0049, points: 0, grade: false, solution: true -->
```python
proba_surv_child = 0 # SOLUTION
# BEGIN SOLUTION
data_child = data[['Survived','Age','Sex']]
data_child = data_child[data_child['Age']<18]

proba_surv_child = (data_child['Survived']==1).mean()
# print("Number of children that survived:", (data_child['Survived']==1).sum())
# print("Number of children that didn't survive:",(data_child['Survived']==0).sum())

data_child_female = data_child[data_child['Sex']=='female']
data_child_male = data_child[data_child['Sex']=='male']
#print(len(data_child_male),len(data_child_female))
# END SOLUTION 
print(proba_surv_child)
```

```python
## test ##
assert_almost_equal(proba_surv_child , 0.5398230088495575 , places=2)
```