```
BEGIN ASSIGNMENT
requirements: requirements.txt
init_cell: true
export_cell: true
files:
    - billboard_intermediate.csv
    - population_data.csv
    - height_2022.csv
generate:
    token: ''
    course_id: 62982        # required
    assignment_id: 2    # required
    filtering: true        # true is the default
    show_stdout: true
    show_hidden: true
    show_all_public: true
    visibility: visible
```

```python
import numpy as np
from pandas import DataFrame as df
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
import otter
grader = otter.Notebook()
```

# Exercise sheet 2
**Hello everyone!**

This is the first graded exercise sheet. <br>
**Points: 15 **

Topics of this exercise sheet are:
* Work with Jupyter
* See autograding
* Lecture 2: Tidy data, pandas

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [ADAM](https://adam.unibas.ch/goto_adam_crs_1266890.html), [Piazza](https://piazza.com/unibas.ch/spring2023/63982), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids23).

Your notebook will be automatically graded using gradescope. Before the exercise session, you should have received an email saying that you were registered to a course in gradescope. If this is not the case, please contact us.

The sheet is released before the exercise session on Tuesday. You have until the next exercise session on Tuesday to submit this sheet. Most of the questions of the exercise sheet can be handle during the session.

**Handout date**: 2023/03/7 <br>
**Submission date**: 2023/03/14 (before 4pm)<br>
**Discussion date**: 2023/03/14

## 1 - Tidy data (9 points)

### A (2 points)
We will use data published by the [World Bank](https://data.worldbank.org/indicator/SP.POP.TOTL) about population in countries.

```python
import pandas as pd

population = pd.read_csv("population_data.csv", skiprows=3).dropna(how='all', axis=1)
population.head()
```

Output:
```
Country Name Country Code     Indicator Name Indicator Code  \
0                        Aruba          ABW  Population, total    SP.POP.TOTL   
1  Africa Eastern and Southern          AFE  Population, total    SP.POP.TOTL   
2                  Afghanistan          AFG  Population, total    SP.POP.TOTL   
3   Africa Western and Central          AFW  Population, total    SP.POP.TOTL   
4                       Angola          AGO  Population, total    SP.POP.TOTL   

          1960         1961         1962         1963         1964  \
0      54208.0      55434.0      56234.0      56699.0      57029.0   
1  130836765.0  134159786.0  137614644.0  141202036.0  144920186.0   
2    8996967.0    9169406.0    9351442.0    9543200.0    9744772.0   
3   96396419.0   98407221.0  100506960.0  102691339.0  104953470.0   
4    5454938.0    5531451.0    5608499.0    5679409.0    5734995.0   

          1965  ...         2011         2012         2013         2014  \
0      57357.0  ...     102050.0   
```

**Question 1**:
Manipulate the dataframe "population" to be tidy. Currently, column headers are values, not variable names.

* Make a new dataframe "population_melted" that contains the same data as "population", but
    * it has a new column "Year" containing the year of the observation
    * it has a new column "Population" containing the number of inhabitants of the country in this year
    * otherwise, it has exactly the same columns and data as "population"
*Hint*: you can use 'pd.melt'.

```
BEGIN QUESTION
name: Question 1        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

```python
population_melted = pd.melt(frame=population, # SOLUTION
# BEGIN SOLUTION
                            id_vars=["Country Name", "Country Code", "Indicator Name", "Indicator Code"], 
                            value_vars=population.columns[4:], # melt from 4th column (so only years)
                            var_name="Year", 
                            value_name="Population")
# END SOLUTION
```

<!-- grade_id: cell-8425c8e5ab3ddda3, points: 0, grade: false, solution: true -->
```python
## test ##
assert_equal(len(population_melted['Year']), 16226)
```

```python
## test ##
assert_equal(len(population_melted['Country Name']),16226)
```

```python
## test ##
assert_is_instance(population_melted,pd.DataFrame)
```

```python
## test ##
tmp = population_melted[(population_melted["Country Code"]=='ABW') & (population_melted["Year"]=='1965')]
assert_equal(np.squeeze(tmp['Population'].values), 57357)
```

```python
## test ##
assert_equal(population_melted.loc[(population_melted["Country Name"] == "Switzerland") & (population_melted["Year"] == "1960")]["Population"].values[0], 5327827.0)
```

**Question 2**: Sort the dataframe population_melted by "Country Name" and "Year"

```
BEGIN QUESTION
name: Question 2        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

```python
population_melted_sort = population_melted.sort_values(by=["Country Name", "Year"]) # SOLUTION
```

```python
## test ##
assert_true(df.equals(population_melted_sort,pd.melt(frame=population, 
                            id_vars=["Country Name", "Country Code", "Indicator Name", "Indicator Code"], 
                            value_vars=population.columns[4:], # melt from 4th column (so only years)
                            var_name="Year", 
                            value_name="Population").sort_values(by=["Country Name", "Year"])))
```

### B (3 points)
We will use data about the average height in countries around the world. It comes from [Kaggle](https://www.kaggle.com/majyhain/height-of-male-and-female-by-country-2022).

```python
height = pd.read_csv("height_2022.csv")
height.head(-1)
```

Output:
```
Rank            Country Name  Male Height in Cm  Female Height in Cm  \
0       1             Netherlands             183.78               170.36   
1       2              Montenegro             183.30               169.96   
2       3                 Estonia             182.79               168.66   
3       4  Bosnia and Herzegovina             182.47               167.47   
4       5                 Iceland             182.10               168.91   
..    ...                     ...                ...                  ...   
193   194               Guatemala             164.36               150.91   
194   195              Mozambique             164.30               155.42   
195   196        Papua New Guinea             163.10               156.89   
196   197         Solomon Islands             163.07               156.79   
197   198                    Laos             162.78               153.10   

     Male Height in Ft  Female Height in Ft  
0                 6.03            
```

**Question 3**: Make a new dataframe "height_tidy" where you removed the columns "Male Height in Ft", "Female Height in Ft"

```
BEGIN QUESTION
name: Question 3        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

```python
height_tidy = height.drop(["Male Height in Ft", "Female Height in Ft"], axis=1) # SOLUTION
```

```python
## test ##
if 'Male Height in Ft' in height_tidy:
    assert_true(False)
```

```python
## test ##
if 'Female Height in Ft' in height_tidy:
    assert_true(False)
```

```python
height_tidy.head()
```

Output:
```
Rank            Country Name  Male Height in Cm  Female Height in Cm
0     1             Netherlands             183.78               170.36
1     2              Montenegro             183.30               169.96
2     3                 Estonia             182.79               168.66
3     4  Bosnia and Herzegovina             182.47               167.47
4     5                 Iceland             182.10               168.91
```

**Question 4**: Manipulate the dataframe "height_tidy" to be tidy. Currently, multiple variables are stored in one column.

* Make a new dataframe "height_melt" that contains the same data as "height_tidy", but
    * it has a new column "Sex" containing a single character "M" or "F" for "male"/"female"
    * it has a new column "Height in Cm" containing data about the height in cm

```
BEGIN QUESTION
name: Question 4        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-b774d9d556851043, points: 0, grade: false, solution: true -->
```python
height_melt = pd.melt(frame=height_tidy, # SOLUTION
# BEGIN SOLUTION
                      id_vars=["Rank", "Country Name"], 
                      value_vars=height_tidy.columns[2:], # melt the height data in cm
                      var_name="Sex", 
                      value_name="Height in Cm") 
# END SOLUTION

# shrink the data on gender to one letter
height_melt["Sex"] = height_melt["Sex"].str[0] # SOLUTION

# show that it works
height_melt.head(5)
```

Output:
```
Rank            Country Name Sex  Height in Cm
0     1             Netherlands   M        183.78
1     2              Montenegro   M        183.30
2     3                 Estonia   M        182.79
3     4  Bosnia and Herzegovina   M        182.47
4     5                 Iceland   M        182.10
```

```python
## test ##
assert_true(isinstance(height_melt,pd.DataFrame))
```

```python
## test ##
assert_true(len(height_melt)== 398)
```

```python
## test ##
assert_true(len(height_melt[height_melt["Sex"] == "M"]) == len(height_melt[height_melt["Sex"] == "F"]))
```

```python
## test ##
assert_true(height_melt.loc[(height_melt["Country Name"] == "Switzerland") & (height_melt["Sex"] == "F")]["Height in Cm"].values[0] == 164.33)
```

```python
## test ##
assert_true(height_melt.loc[(height_melt["Country Name"] == "Switzerland") & (height_melt["Sex"] == "M")]["Height in Cm"].values[0] == 178.73)
```

### C (4 points)
We will work with data about billboards in the 2000s (as in the lecture).

```python
billboard = pd.read_csv("billboard_intermediate.csv")
billboard.head()
```

Output:
```
year artist.inverted                       track  time genre        date  \
0  2000           Nelly  (Hot S**t) Country Grammar  4:17   Rap  2000-04-29   
1  2000           Nelly  (Hot S**t) Country Grammar  4:17   Rap  2000-05-06   
2  2000           Nelly  (Hot S**t) Country Grammar  4:17   Rap  2000-05-13   
3  2000           Nelly  (Hot S**t) Country Grammar  4:17   Rap  2000-05-20   
4  2000           Nelly  (Hot S**t) Country Grammar  4:17   Rap  2000-05-27   

  date.peaked  week   rank  
0  2000-09-16     1  100.0  
1  2000-09-16     2   99.0  
2  2000-09-16     3   96.0  
3  2000-09-16     4   76.0  
4  2000-09-16     5   55.0
```

**Question 5:**: Manipulate the dataframe "billboard" into two dataframes that each store one observational unit.

* Make a new dataframe "billboard_songs" that contains information about the songs
    * It should contain these columns from the "billboard" dataframe: "artist.inverted", "track", "time", "genre"
* Make a new dataframe "billboard_charts" that contains information about the chart performance of the songs
    * It should contain these columns from the "billboard" dataframe: "date", "date.peaked", "week", "rank"
* Additionally, both dataframes should contain a column named "track_id" that gives each track a unique ID and links the information in the two dataframes.

(Hint: Pandas has a function [.groupby().ngroup()](https://stackoverflow.com/questions/39665374/how-to-make-a-group-id-using-pandas) that could help you with the track id)

```
BEGIN QUESTION
name: Question 5        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 4      # how many points this question is worth; defaults to 1 internally
check_cell: true  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-262fb79928ed7861, points: 0, grade: false, solution: true -->
```python
billboard["track_id"] = billboard.groupby(["artist.inverted", "track", "time", "genre"]).ngroup() # SOLUTION

# BEGIN SOLUTION
billboard_songs = billboard[["track_id", "artist.inverted", "track", "time", "genre"]]
billboard_songs = billboard_songs.drop_duplicates()
# END SOLUTION
billboard_songs = billboard_songs.reset_index() # SOLUTION
billboard_songs.head()

billboard_charts = billboard[["track_id", "date", "date.peaked", "week", "rank"]] # SOLUTION
billboard_charts.head()
```

Output:
```
track_id        date date.peaked  week   rank
0       222  2000-04-29  2000-09-16     1  100.0
1       222  2000-05-06  2000-09-16     2   99.0
2       222  2000-05-13  2000-09-16     3   96.0
3       222  2000-05-20  2000-09-16     4   76.0
4       222  2000-05-27  2000-09-16     5   55.0
```

```python
## test ##
assert_true(isinstance(billboard_songs, pd.DataFrame))
```

```python
## test ##
assert_true(isinstance(billboard_charts, pd.DataFrame))
```

```python
## test ##
assert_true(np.array_equal(np.unique(billboard_songs["track_id"]),np.unique(billboard_charts["track_id"])))
```

```python
## test ##
assert_true(len(billboard_songs) == 317)
```

```python
## test ##
assert_true(len(billboard_charts) == 24092)
```

## 2 - EDA (6 points)

### A (2 points)
Let's answer one more question about billboard hits in the 2000s. Use the billboard, billboard_songs or billboard_charts dataframes.

**Question 6**: What was the number one hit (rank = 1) on 2000/07/29? This question may seem random (and it is), but this is the 2000s! This is the day on which Jennifer Aniston married Brad Pitt. Did they play this song? Who knows.

* Define one variable
    * "name_song": Name of the song that was at rank 1 on the given date

```
BEGIN QUESTION
name: Question 6        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-75a4834772643fcd, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
id_song = billboard_charts.loc[(billboard_charts["rank"] == 1) & (billboard_charts["date"] == "2000-07-29")]
# END SOLUTION
name_song = billboard_songs[billboard_songs["track_id"] == id_song["track_id"].values[0]]["track"].values[0] # SOLUTION
name_song
```

Output:
```
"It's Gonna Be Me"
```

<!-- grade_id: cell-03b8dca8c8ea670c, points: 2, grade: true, solution: false -->
```python
## test ##
assert_true(name_song == "It's Gonna Be Me")
```

### 2b (4 points)
Let's answer some questions about average heights. Use the height, height_tidy or height_melt dataframes.

**Question 7**: What's the country with the tallest men?

* Define two variables
    * "country_tallest_men": Name of the country of tallest men as a string
    * "height_tallest_men": Average height of men in this country in cm

(Hint: Pandas function max() can help here!)

```
BEGIN QUESTION
name: Question 7        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-4283fe78120f5a5d, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
tallest_men = height_melt.loc[height_melt[height_melt["Sex"] == "M"]["Height in Cm"].idxmax()] 
# END SOLUTION

country_tallest_men = tallest_men["Country Name"] # SOLUTION
height_tallest_men = tallest_men["Height in Cm"] # SOLUTION
```

```python
## test ##
assert_true(country_tallest_men == "Netherlands")
```

```python
## test ##
assert_almost_equal(height_tallest_men, 183.78)
```

**Question8**: In which country is the difference of average heights between men and women largest?

* Define two variables
    * "country_largest_difference": Name of the country where the average height difference between men and women is largest
    * "largest_difference": Difference of average heights of men and women in this country in cm

```
BEGIN QUESTION
name: Question 8        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 2      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-4e575ca703ccbd4e, points: 0, grade: false, solution: true -->
```python
height["diff heights"] = height["Male Height in Cm"] - height["Female Height in Cm"] # SOLUTION

# BEGIN SOLUTION
height.iloc[height["diff heights"].idxmax()]
# END SOLUTION

country_largest_difference = height.iloc[height["diff heights"].idxmax()]["Country Name"] # SOLUTION
largest_difference = height.iloc[height["diff heights"].idxmax()]["diff heights"] # SOLUTION
```

<!-- grade_id: cell-f94f6cefaae9dc5b, points: 2, grade: true, solution: false -->
```python
## test ##
assert_true(country_largest_difference == "Puerto Rico")
```

```python
## test ##
assert_almost_equal(largest_difference, 16.42)
```

**Question 9**: What's (Pearson's) correlation value of average male and average female height?

* Define a new variable "corr_avg_heights" that contains the correlation value of average male and female heights

(Hint: Use Panda's function corr().)

```
BEGIN QUESTION
name: Question 9        # (required) the path to a requirements.txt file
manual: false     # whether this is a manually-graded question
points: 1      # how many points this question is worth; defaults to 1 internally
check_cell: false  # whether to include a check cell after this question (for autograded questions only)
```

<!-- grade_id: cell-b5438d2320f5d36d, points: 0, grade: false, solution: true -->
```python
corr_avg_heights = height.corr()["Male Height in Cm"]["Female Height in Cm"] # SOLUTION
```

Output:
```
/tmp/ipykernel_48470/85255859.py:1: FutureWarning: The default value of numeric_only in DataFrame.corr is deprecated. In a future version, it will default to False. Select only valid columns or specify the value of numeric_only to silence this warning.
  corr_avg_heights = height.corr()["Male Height in Cm"]["Female Height in Cm"] # SOLUTION
```

<!-- grade_id: cell-21e6f8a0134468cf, points: 1, grade: true, solution: false -->
```python
## test ##
assert_almost_equal(corr_avg_heights,0.9287870871294992)
```

Output:
```
Cell In[38], line 2
      1 ## Hidden test ##
----> 2 assert_almost_equal(corr_avg_heights,0.9287870871294992)

File ~/anaconda3/lib/python3.9/unittest/case.py:868, in TestCase.assertAlmostEqual(self, first, second, places, msg, delta)
    865 if delta is not None and places is not None:
    866     raise TypeError("specify delta or places not both")
--> 868 diff = abs(first - second)
    869 if delta is not None:
    870     if diff <= delta:

TypeError: unsupported operand type(s) for -: 'str' and 'float'
```